<a href="https://colab.research.google.com/github/hallque/Fall2022/blob/main/Assignment1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 1

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. 

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Runtime, and select restart runtime and run all cells.

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work. When you click on the link in Canvas you open a shared notebook that you cannot save. Make sure to copy the notebook to your own Google Drive before starting work. 

5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.

----

```python
%%capture
#@title Run this cell before getting started
!pip install nose
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import os, sys
import math
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

----
# Problem 1: Creating a Python string

- Type "x = 'This is my first string. Hookem Horns' " without the outer quotation marks in the below cell and run it.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
#@title This is a grading cell. Run after completing the cell above.
try:
    if ((x=='This is my first string. Hookem Horns') & (type(x)==str)):
        print('Great! You created a string!')
    else:
        print('x is not the right value check your code and run again')
except NameError:
    print('x is not defined! You need to run the code cell!')
    
assert_equal(type(x),str, msg='x is not a string!')
```

----
# Problem 2: Import numpy

- Sometimes we want to do things in Python that are not supported natively. Luckily Python supports thousands of packages (e.g., numpy, pandas, etc.) that increase functionality. As we will discuss in more detail later on in the course, these packages are easy to install as needed. For now, type "`import numpy`" in the below cell to import the numpy module.

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
#@title This is a grading cell. Run after completing the cell above.
if 'numpy' in dir():
    print('numpy module has been successfully imported')
else:
    print('numpy module has not been imported')
    
assert_equal(numpy.array([1])[0],1,msg='Make sure you ran the code cell to import numpy!')
```

----
# Problem 3: Add emphasis to Markdown text

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that does the following

    - Writes the word "bold" **in bold**
    
    - Writes the word "italics" *in italics*
    
    - Write the phrase "scrach and bold" as strikethrough _and_ bold, like so ~~**scratch and bold**~~. The symbols that initialize the bold enviroment should go on the inside. 
    
    

- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the bold/italicize/scratch environments, and your text within those environments.

- HINT: We have not learned how to render strikethrough. Can you think of a way to figure out how to do it without searching the internet? 

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
#@title This is a grading cell. Run after completing the cell above.
text1 = md.find('~**')
assert_not_equal(text1,-1,msg='Make sure you are using scratch first and bold second.')
text2 = md.find('**')
assert_not_equal(text2,-1,msg='You did not bold the word bold.')
text3 = md.find('_ital')
assert_not_equal(text3,-1,msg='You did not italicize the word italics')

if text1 != -1 & text2 != -1 & text3 != -1:
  print('Good work!')
else:
  print('')
```

----
# Problem 4: Create a Markdown list

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that creates a markdown list with 3 items. The items should be numbered 1 through 3, respectively, and there should be a `.` character after each number. Name each item "Item 1", "Item 2", and "Item 3", respectively. Note that the "I" in Item sould be capitalized.

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
#@title This is a grading cell. Run after completing the cell above.
text1 = md.find('1.')
assert_not_equal(text1,-1,msg='Number 1. is mising from your markdown list')
text2 = md.find('2')
assert_not_equal(text2,-1,msg='Number 2. is mising from your markdown list')
text3 = md.find('3.')
assert_not_equal(text3,-1,msg='Number 3. is mising from your markdown list')
text4 = md.find('Item')
assert_not_equal(text4,-1,msg='Make sure you named each item in the list properly.')

if text1 != -1 & text2 != -1 & text3 != -1 & text4 != -1:
  print('Good work!')
else:
  print('')
```

----
# Problem 5: Show code in Markdown

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that renders Python code. The Python code to be rendered is the statement `x = 'this is a python string'`. The rendered code should look just like it does in this example.


- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the "code" environment.

- HINT: You can edit this cell to see how it is rendering the statement as code.

<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
md='''

### YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
#@title This is a grading cell. Run after completing the cell above.
text1 = md.find('`x')
assert_not_equal(text1,-1,msg='Make sure you are using the right characters to render the code.')
text2 = md.find('python')
assert_not_equal(text2,-1,msg='Make sure you type the string correctly.')
if text1 != -1 & text2 != -1:
  print('Good work!')
else:
  print('')
```